# Welcome to Lab_Plots! 🕵🏻‍♂️ 📊 🕵🏻‍♀️

Fun fact: Every employee at the University of Illinois is a "public employee" and all public employee salaries are [publicly available online](https://www.bot.uillinois.edu/resources/gray_book) -- we have curated this data in a cleaned dataset for you to explore!  This includes data about every professor, administrator, and football coach!

The goal of this lab is to work with **real UIUC salary data** to explore its properties, answer important questions, and to think about the implications of collecting and analyzing this data.  Throughout the lab, it is important to think about being a critical consumer of data who can not only use statistics and programming to analyze data, but can also think about the **"why"** part of data science both in the classroom and in the world. Let’s get started!

In this lab, you will:
- Work with real UIUC salary data to explore some of the statistics that we talked about in lecture: mean, median, standard deviation, etc.
- Practice creating plots to **visualize quantitative data**: boxplots and histograms.
- See how data science can be used in the real world to think about important issues through written individual reflections and discussions with your group.

A few tips to remember:

- **You are not alone on your journey in learning programming!**  You have your lab TA, the CAs, your lab group, and the professors (Prof. Wade and Prof. Karle), who are all here to help you out!
- If you find yourself stuck for more than a few minutes, ask a neighbor or course staff for help!  When you are giving help to your neighbor, explain the **idea and approach** to the problem without sharing the answer itself so they can have the same **<i>ah-hah</i>** moment!
- We are here to help you!  Don't feel embarrassed or shy to ask us for help!

Let's get started!


<hr style="color: #DD3403;">

```python
# Meet your CAs and TA if you haven't already!
# First name is enough, we'll know who they are! :)
ta_name = ""
ca1_name = ""
ca2_name = ""
ca3_name = ""

# Work with your group again this week! 
#
# QOTD to Ask Your Group: "What was your favorite childhood game?"
partner1_name = ""
partner1_netid = ""
partner1_fav_game = ""

partner2_name = ""
partner2_netid = ""
partner2_fav_game = ""

partner3_name = ""
partner3_netid = ""
partner3_fav_game = ""

# Which lab section are you in? Record it below! It should start with a Y and have 
# 2 numbers after the Y. If you're unsure - ask your TA/CAs!
lab_section = "Y"
```

<hr style="color: #DD3403;">

## Setup: Import the Graybook Dataset

The "Gray Book" is historical term for the book of "Academic and Administrative Appointments".  As a public university, all positions (including job title, tenure status, and salary) at UIUC are publicly approved by the Board of Trustees.  After approval, they are published publicly at [https://www.bot.uillinois.edu/resources/gray_book](https://www.bot.uillinois.edu/resources/gray_book).

We have parsed the HTML tables and done a little data cleaning for you. The "Graybook Dataset" provided here includes all faculty (except for the Division of Collegiate Athletics, for salary outlier reasons) at the University of Illinois, based on the **2022-2023 Graybook report**.  A CSV version of this dataset is available at the following URL:

```
https://waf.cs.illinois.edu/discovery/graybook.csv
```

Import `pandas` and load this dataset into a DataFrame, `df`:

<!-- grade_id: cell-27971beecac131cc, points: 0, grade: false, solution: true -->
```python
...
```

### 🔬 Test Case Checkpoint 🔬

```python
## == TEST CASES for Loading in Graybook ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with a message (with the emoji showing), you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.
assert( 'df' in vars()), "Load the dataset into the variable named `df`."
assert ( len(df) == 6288 ), "This is not the Graybook dataset you're looking for."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")   
```

<hr style="color: #DD3403;">

## Part 1: Exploratory Data Analysis (EDA)

As discussed in lecture, the first step of any data analysis is to **get familiar** with your dataset.  Think about what this data can tell you and what variables are included.  Data scientists always start with this step.

Let’s do some general **exploratory data analysis** to feel out our dataset.  Before you do any calculations, ponder this question:

### Initial Guess

**What do you estimate the average salary of all UIUC Faculty to be?**

*(✏️ Edit this cell to replace this text with your answer. ✏️)*

### Puzzle 1.1: Descriptive Statistics

Our Graybook Dataset contains both the `Present Salary` and `Proposed Salary` of employees at U of I. For now, we're only interested in the `Present Salary` column.

Using `df`, find the following information:

1. The number of faculty at UIUC, stored in the variable `num_employee` (Hint: each row is an employee!)
2. The **mean** present salary, storing in the variable `mean_sal`
3. The **median** present salary, storing in the variable `median_sal`
4. The **standard deviation** of present salary, storing in the variable `std_sal`

<!-- grade_id: cell-419796a0b415e588, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: cell-4dbbeef1022a20e4, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: cell-727195c3fef2cf78, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: cell-35e651d14a628e50, points: 0, grade: false, solution: true -->
```python

```

### 🔬 Test Case Checkpoint 🔬

<!-- grade_id: cell-04e624817d681f8e, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with a message (with the emoji showing), you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.
import math

x = mean_sal + median_sal + std_sal
assert( num_employee == 6288 ), "Your calculations of the number of employees is incorrect."
assert( math.isclose(x, 303895.7881426731) ), "Your calculations of mean, median, or standard deviation is incorrect."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### Visual Displays of Data: A Key Part of EDA!

Now, we are a bit more familiar with the dataset through summary statistics. Looking at overall descriptive statistics helps us summarize all of the observations in a column, rather than having to scroll through all of the observations!  However, descriptive statistics alone often don’t tell the whole story. This is where having tools for **visualizing statistics** comes in handy.

### Puzzle 1.2: Boxplots and Histograms

Next, let’s look at two simple, yet powerful visualizations: the **boxplot** and the **histogram**! 

Generate a **boxplot** of the `Present Salary` column.


```python
...
```

Generate a **histogram** of the `Present Salary` column.

```python
...
```

###  Puzzle 1.3: Remove the Outliers

These visualizations that you just created look different than any that you saw in lecture! In class we learned how to calculate outliers mathematically using Q1, Q3, and the IQR. Create a new DataFrame, `df_without_outliers`, by removing all employees whose `Present Salary` is an outlier.

**Hint**: You can find Q1 and Q3 in Python using the `.quantile()` function. Once you have these numbers, look in your notebook to see the conditions for high and low outliers and write a complex conditional to select the rows where the `Present Salary` is not an outlier.

```python
# Use this cell to calculate Q1, Q3, IQR
```

```python
# Use this cell to calculate the cutoffs for the outliers
```

```python

```

### 🔬 Test Case Checkpoint 🔬

```python
## == TEST CASES for filtering the data to include non-outliers ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with a message (with the emoji showing), you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.
assert( 'df_without_outliers' in vars()), "Load the dataset into the variable named `df_without_outliers`."
assert ( len(df_without_outliers) == 5984 ), "You did not filter the dataset correctly."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")   
```

### Part 1.4: Boxplots (Round 2)

Using `df_without_outliers`, create a boxplot of the `Present Salary` for the employees whose salary is not an outlier.

```python
...
```

Note: You can see that the boxplot created above still has outliers. This is because the boxplot is created from the new DataFrame `df_without_outliers` rather than the original DataFrame `df`, calculating new outliers from our `df_without_outliers` data.

### Part 1.5: Histograms (Round 2)

Using `df_without_outliers`, create a histogram of the `Present Salary` for the employees whose salary is not an outlier.

```python
...
```

<hr style="color: #DD3403;">

## Part 2: Department-Level Data

Our Graybook dataset also includes the `Department Name` of each employee. This means you can see what department each professor is from and look at other employees in the same department.  

We can calculate **descriptive statistics** and look at visual displays of data on a **department level**!  Let’s explore whether or not there are salary **discrepancies** among departments.

### Puzzle 2.1: Grouping the Data

In lecture, you learned syntax that will help you **group** all of the people in each department so you can find departmental aggregates.  In other words, we want to group by `'Department Name'`.

To make things easier to read, we will first modify your DataFrame to include only the `Department Name` and `Present Salary` fields before you `groupby`.

First, using the cell below, create a second DataFrame named `depts` **containing only** the `Department Name` and `Present Salary` columns. Remember to use **double square-brackets** to specify these columns. Here's an example of the syntax: `df[['col1', 'col2']]`.
- Make sure to use `df` that contains ALL employees.

```python

```

Now, find the mean Present Salary for each department to create `agg_mean_df`:

<!-- grade_id: cell-72189b308296916b, points: 0, grade: false, solution: true -->
```python

```

Hint: Remember to `reset_index()`

### 🔬 Test Case Checkpoint 🔬

<!-- grade_id: cell-d4a69c902629304a, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2.1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with a message (with the emoji showing), you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.
import math
assert ( len(depts.columns) == 2 ), "Make sure your depts contains only the Department Name and Present Salary columns."
assert ( len(agg_mean_df.columns) == 2), "Make sure your agg_mean_df contains only the Department Name and Present Salary columns."
assert ("agg_mean_df" in vars()), "Make sure your aggregated DataFrame of means is named 'agg_mean_df'."
assert( len(agg_mean_df) == 258 ), "You should have 258 different departments.  Check to make sure you've grouped the data correctly."
assert( "Present Salary" in agg_mean_df ), "Your agg_mean_df is missing the Salary column."
assert( math.isclose(agg_mean_df['Present Salary'].mean(), 103214.15081569373 ) ), "Your salary values are incorrect. Double-check that you are aggregating by `mean`."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### Puzzle 2.2: Department Comparisons

Let’s look at how the mean salaries differ across departments.

Using your `agg_mean_df`, create a barplot for the departments with the top-40 average salaries. Make sure to change the figsize so that the plot looks better and that your x-axis is labeled appropriately.

```python

```

### Puzzle 2.3: Grouping the Data, Again

We have an aggregated DataFrame of the salaries by department, but what if we want to find the number of employees in each department? 

Each row represents a single employee, we can use this to find the number of employees per department. In the cell below, using your `depts` from earlier, create `agg_count_df` which finds the number of employees in each department:

```python

```

If done correctly, your `agg_count_df` should now contain the number of employees in each department under the `Present Salary` column. But this column name is **misleading**. 

**Run the provided cell below** to rename the `Present Salary` column to `Num Employees`:

```python
agg_count_df.rename(columns={"Present Salary": "Num Employees"}, inplace=True)
agg_count_df
```

### Puzzle 2.4: Departments With The Most Employees!

Use the cell below to find the 10 departments that have the most employees.

```python
...
```

### 🔬 Test Case Checkpoint 🔬

```python
## == TEST CASES for Puzzle 2.3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with a message (with the emoji showing), you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.
import math
assert ( len(agg_count_df.columns) == 2 ), "Make sure your depts contains only the 2 columns."
assert ("agg_count_df" in vars()), "Make sure your aggregated DataFrame of employee counts is named 'agg_count_df'."
assert( len(agg_count_df) == 258 ), "You should have 258 different departments.  Check to make sure you've grouped the data correctly."
assert( "Num Employees" in agg_count_df ), "Your agg_count_df is missing the Num Employees column. Make sure you ran the provided cell to rename the column."
assert( math.isclose( agg_count_df['Num Employees'].mean(), 24.372093023255815 ) ), "Your employee counts are incorrect. Double-check that you are aggregating by `count`."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### Puzzle 2.5: Salaries and Number of Employees

To visualize the relationship between the Departments' number of employees and their average salaries, we are going to create a scatter plot. In order to create the scatter plot, we need to use the appropriate columns from `agg_mean_df` and `agg_count_df`:

In the DataFrame `agg_count_df`, create a new column called "Average Salary" which contains the average salary for each department. (Remember: You have already calculated the average salary for each department in `agg_mean_df`. You can use that column)

Then create a scatterplot using x = "Num Employees" and y = "Average Salary". Make sure you add a title to your scatterplot!

```python
# Create a new column in agg_count_df
```

```python
# Create your scatterplot
```

<hr style="color: #DD3403;">

## Part 3: Gender and Salaries

Data can often reveal **systemic problems or discrimination**. For example, in many companies, men and women are promoted at **different rates**.  Let’s look at a subset of the salary dataset to investigate whether or not there is a **difference in salaries** between faculty who identify as men and women in two departments: **STAT** and **CS** (Karle and Wade’s home departments). 

We've compiled data from these departments, added a `Gender` column, and placed it in a dataset called `STAT_CS_gender.csv`. 

This data is also located in the **same directory as this lab**. To load it in, just specify the **local file path** (`"STAT_CS_gender.csv"`)!

### Puzzle 3.1: Loading Our Dataset
Using the cell below, read `STAT_CS_gender.csv`, store it in a variable called `STAT_CS_df`, and display it to see what it looks like!

<!-- grade_id: cell-dca5763613817466, points: 0, grade: false, solution: true -->
```python

```

Now, let's create **two subsets** of our `STAT_CS_df`. 

Using conditionals in the cells below, create:
- `STAT_CS_M`, a `DataFrame` of the staff and faculty who identify as Male (**"M"**) under the `Gender` column 
- `STAT_CS_F`, a `DataFrame` of the staff and faculty who identified as Female (**"F"**) under the `Gender` column 

```python

```

```python

```

### 🔬 Test Case Checkpoint 🔬

```python
## == TEST CASES for Puzzle 3.1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any errors, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.
assert("STAT_CS_df" in vars()), "Ensure you've named your original DataFrame `STAT_CS_df`."
assert(len(STAT_CS_df) == 145), "This is not the STAT_CS_df you are looking for."
assert("STAT_CS_M" in vars()), "Ensure your male subset of STAT_CS_df is named `STAT_CS_M`."
assert(len(STAT_CS_M) == 106), "Double check your conditional to generate STAT_CS_M - the number of rows is incorrect."
assert("STAT_CS_F" in vars()), "Ensure your female subset of STAT_CS_df is named `STAT_CS_F`."
assert(len(STAT_CS_F) == 39), "Double check your conditional to generate STAT_CS_F - the number of rows is incorrect."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### Puzzle 3.2: Merging Two Columns into One DataFrame

To create a new DataFrame containing only the data you found above, the following provided line of code creates `df_salary_by_gender`:

```python
df_salary_by_gender = pd.DataFrame({
    "female": STAT_CS_F["Present Salary"],
    "male": STAT_CS_M["Present Salary"],
})
```

Take a look at the DataFrame `df_salary_by_gender`.  Every row will either have data for in the `female` column or the `male` column:

```python
df_salary_by_gender
```

### Puzzle 3.3: Visualization

Let's create a visualization! Using the next cell, create a **boxplot** of `df_salary_by_gender` that includes both genders:

```python
...
```

### Puzzle 3.4: EDA, Again

Now that we've visualized the data, let's explore some basic statistics once more to gain further insight. 

In the following cells, calculate:
- The **mean** `Present Salary` for **Male** STAT/CS Faculty, storing in the variable `mean_m`
- The **median** `Present Salary` for **Male** STAT/CS Faculty, storing in the variable `median_m`
- The **standard deviation** of `Present Salary` for **Male** STAT/CS Faculty, storing in the variable `std_m`

```python

```

```python

```

```python

```

Now, in the cells below, calculate:

- The **mean** `Present Salary` for **Female** STAT/CS Faculty, storing in the variable `mean_f`
- The **median** `Present Salary` for **Female** STAT/CS Faculty, storing in the variable `median_f` 
- The **standard deviation** of `Present Salary` for **Female** STAT/CS Faculty, storing in the variable `std_f`

```python

```

```python

```

```python

```

Run the following cell to make a summary table of your previously calculated data:

```python
pd.DataFrame([
  {"Gender": "F", "Mean ($)": round(mean_f), "Median ($)": round(median_f), "Standard Deviation ($)": round(std_f)},
  {"Gender": "M", "Mean ($)": round(mean_m), "Median ($)": round(median_m), "Standard Deviation ($)": round(std_m)},
])
```

### 🔬 Test Case Checkpoint 🔬

```python
## == TEST CASES for Puzzle 3.4 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with a message (with the emoji showing), you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.
import math
assert( math.isclose( mean_f, 157847.477949 ) )
assert( math.isclose( median_f, 138000.0 ) )
assert( math.isclose( std_f, 50583.799487 ) )
assert( math.isclose( mean_m, 157679.508208 ) )
assert( math.isclose( median_m, 152500.0 ) )
assert( math.isclose( std_m, 49368.327380 ) )

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

<hr style="color: #DD3403;">

## Part 4: Exploring Your Own Interests


At this point of the lab, we have investigated a lot of questions.

However, these have been questions that **we told you to answer**. As a data scientist, it is important to be able to use the data science skills that you learn in the classroom to answer questions that **you have**.

Think about **two questions** that you have about the **Graybook** or **STAT_CS_Gender** datasets that have not been answered. These can be simple questions. Record them below.  Then, answer at least one of these questions using Python and either dataset.  

*(✏️ Edit this cell to replace this text with your two questions. ✏️)*


Now, use the cell below to **find the answer** to **one** of your questions! Remember, it can be something simple. 

```python
...
```

### Analysis: lab_plots takeaways

Please answer the following questions in complete sentences:

**Question 1:** In this lab, you explored two types of data visualizations: box plots and histograms. Explain why is it valuable to look at visual displays of salary data in general (as opposed to just looking at descriptive statistics like the mean and standard deviation)?

**Question 2:** In Part 3, you calculated descriptive statistics of the `Present Salary` of Male and Female STAT/CS Faculty. Based on the work you did in this section, can we draw any conclusions about gender-based salary discrimination from our data? Use the analysis you did to support your answer.

**Question 3:** What was the most interesting thing you learned from doing this lab?

*(✏️ Write your answers to the analysis here. Make sure to use complete sentences and be as descriptive as possible in your answers.✏️)*

<hr style="color: #DD3403;">

## **Submission** 


You're almost done! All you need to do is to commit your lab to GitHub:

1. Make certain to save your work. To do this, go to **File => Save All**

2. After you have saved, exit this notebook and follow the Canvas instructions to commit this lab to your Git repository!

3. Your TA will grade your submission and provide you feedback after the lab is due. :)